---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: 2bafef99a3
lang: r
xp: 100
skills: 1
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## Insert exercise title here

```yaml
type: NormalExercise
key: 224428e878
xp: 100
```

test if outside pacakges can be installed in datacampe

`@instructions`
test if outside pacakges can be installed in datacampe

`@hint`


`@pre_exercise_code`
```{r}
library(stringr)
library(tidyverse)
```

`@sample_code`
```{r}
a<-5
```

`@solution`
```{r}
5
```

`@sct`
```{r}

```
